Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "SIDDHARTH YASHASWEE"
COLLABORATORS = ""
```

```python
import numpy as np
import openmesh as om
```

# Part A
Write a function that reads `spot/spot_triangulated.obj` file and returns the number of vertices and faces in the mesh as a tuple.

<!-- grade_id: n_verts_and_faces, points: 0, grade: false, solution: true -->
```python
def n_verts_and_faces():
    ### BEGIN SOLUTION
    mesh = om.read_trimesh("spot_triangulated.obj")
    return (mesh.n_vertices(), mesh.n_faces());
    ### END SOLUTION
```

<!-- grade_id: n_verts_and_faces_test, points: 0, grade: true, solution: false -->
```python
assert n_verts_and_faces() == (2930, 5856)
### BEGIN HIDDEN TESTS
assert len(n_verts_and_faces()) == 2
### END HIDDEN TESTS
```