# Distance Measures
In this notebook we will implement the approximate computation of geodesic distances via the heat method. We will also compute the Chamfer distance between two point clouds.

```python
from ipywidgets import interact, fixed
import numpy as np
import scipy as sp
from scipy.sparse import csr_matrix, diags
from scipy.sparse.linalg import spsolve, lsqr
from scipy.spatial.distance import cdist
import openmesh as om
import k3d
```

```python
mesh = om.read_trimesh("spot.obj")
bob_mesh = om.read_trimesh("bob.obj")
A = sp.sparse.load_npz("A.npz")
W = -sp.sparse.load_npz("W.npz")
```

## Heat Method
For the heat method you will have to compute the heat distribution. Then you have to derive gradients for this distribution on the triangle mesh and then solve a Poisson equation to compute a geodesic distance field.

In order to compute the heat distribution with `compute_heat_distribution` you will need the (uninverted) mass matrix `A` and the weak Laplacian Beltrami Operator `W` for the spot mesh (stored in `mesh`). We have precomputed and loaded those matrices for you above. Furthermore, we require a time parameter `t` (in our case `t=1`) and an intial heat distribution for each vertex of the mesh. Here we apply heat to a single point (the vertex with id 1570) for which we want to compute the geodesic distance field. Since `A` and `W` are sparse matrices use `spsolve` to solve the linear equation system. The result should be a numpy array with one axis with scalar values for every vertex of the mesh. The heat distribution should then look like [this](./heat_distribution.html) (both the intial heat distribution and the one at time step `t` are shown).

Next, we need to compute the gradient of the heat distribution on our mesh with `compute_gradient_on_trimesh`. Here `mesh` is the spot mesh and `u` the heat distribution we computed above (per vertex). The result should be a 3D gradient vector per triangle (again as a numpy array). You can use the `compute_triangle_area` and `compute_normal` helper functions per triangle (`a`, `b`, and `c` are the 3 3D points of the triangle). In order to iterate over the vertices of each face you can either use `mesh.fv(fh)` for a particular face handle `fh` or get all face vertex indices with `mesh.fv_indices()` as a numpy array. Afterwards you will have to transform the gradient to a vector field that matches the gradient of a distance field with `compute_gradient_of_distance_field`. This function takes as input a numpy array with 3D gradients per triangle face and returns a vector field as a numpy array with 3D vectors per face. If you implemented everything correctly the resulting vector field should look like [this](./gradient_field.html).

Now that we have the gradient field for our distance field, we need to recover the distance field in question. To this end we have to compute the divergence of that gradient field in order to setup and solve a Poisson equation. Implement the computation of the divergence with `compute_divergence_on_trimesh` that takes as as input the spot mesh `mesh` and the vector field computed above with `compute_gradient_of_distance_field` as `vec`. It returns a numpy array with a scalar per vertex. You can use `mesh.vertices()` to iterate over all vertices of the mesh. Use `mesh.voh(vh)` to iterate over all outgoing halfedges for a vertex handle `vh`. You can get the corresponding (unique) face handle for a halfedge handle `hh` with `mesh.face_handle(hh)`. To jump to the next halfedge handle from a halfedge handle `hh` in a face use `mesh.next_halfedge_handle(hh)`. The vertex handle that a halfedge `hh` points to can be looked up with `mesh.to_vertex_handle(hh)`. The cotanget for two vectors $v_1$ and $v_2$ can be computed as $\frac{v_1^Tv_2}{||v_1 \times v_2||}$. If you implemented everything correctly your divergence should look like [this](./divergence.html). You can now setup and solve the Poisson equation in `compute_distance_field` that takes as input the weak Laplace Beltrami Operator and the divergence field. It returns the distance field (a numpy array with scalars per vertex). Use `lsqr` to solve the Poisson equation since it is over-determined. Your solution should look like [this](./distance_field.html).
Note that the distance field can have negative values, as the solution of the Poisson equation is invariant to a constant offset.

<!-- grade_id: cell-0bdcb31cded6ad87, points: 0, grade: false, solution: true -->
```python
def compute_triangle_area(a, b, c):
    return np.linalg.norm(np.cross(b-a, c-a))/2

def compute_normal(a, b, c):
    n = np.cross(b-a, c-a)
    n /= np.linalg.norm(n)
    return n

def compute_gradient_on_trimesh(mesh, u):
    ### BEGIN SOLUTION
    fv_indices = mesh.fv_indices()
    pts = mesh.points()
    grads = np.zeros((fv_indices.shape[0], 3))
    for i in range(fv_indices.shape[0]):
        face = fv_indices[i]
        face_pts = pts[face]
        face_u = u[face]
        n = compute_normal(face_pts[0], face_pts[1], face_pts[2])
        A_f = compute_triangle_area(face_pts[0], face_pts[1], face_pts[2])
        grad = np.zeros(3)
        for j in range(3):
            e_j = face_pts[(j+1)%3] - face_pts[j%3]
            u_j = face_u[(j+2)%3]
            grad += u_j * np.cross(n, e_j)
        grad /= 2*A_f
        grads[i] = grad
    return grads
    ### END SOLUTION

def compute_divergence_on_trimesh(mesh, vec):
    ### BEGIN SOLUTION
    div = np.zeros(mesh.n_vertices())
    pts = mesh.points()
    for vh in mesh.vertices():
        a = pts[vh.idx()]
        for hh in mesh.voh(vh):
            fh = mesh.face_handle(hh)
            next_hh = mesh.next_halfedge_handle(hh)
            b = pts[mesh.to_vertex_handle(hh).idx()]
            c = pts[mesh.to_vertex_handle(next_hh).idx()]

            e_1 = b - a
            e_2 = c - a
            v_1 = a - c
            v_2 = b - c
            cot_theta_1 = np.dot(v_1, v_2) / np.linalg.norm(np.cross(v_1, v_2))
            v_1 = c - b
            v_2 = a - b
            cot_theta_2 = np.dot(v_1, v_2) / np.linalg.norm(np.cross(v_1, v_2))
            div[vh.idx()] += cot_theta_1*np.dot(e_1, vec[fh.idx()]) + cot_theta_2*np.dot(e_2, vec[fh.idx()])
        div[vh.idx()] *= 0.5
    return div
    ### END SOLUTION
    
def compute_heat_distribution(A, W, t, init_heat):
    ### BEGIN SOLUTION
    u_t = spsolve(A-t*W, A*init_heat)
    # right side does not need to be multiplicated with A as it is nonzero at single point anyways
    return u_t
    ### END SOLUTION

def compute_gradient_of_distance_field(grads):
    ### BEGIN SOLUTION
    return -grads / np.linalg.norm(grads, axis=-1, keepdims=True)
    ### END SOLUTION

def compute_distance_field(W, div):
    ### BEGIN SOLUTION
    return lsqr(W, div)[0]
    ### END SOLUTION
```

<!-- grade_id: cell-f8709265f219aab5, points: 1, grade: true, solution: false -->
```python
init_u = np.zeros(mesh.n_vertices())
init_u[1570] = 1.
t = 1
u_t = compute_heat_distribution(A, W, t, init_u)
np.testing.assert_approx_equal(u_t[0], 0.00012, significant=2)
np.testing.assert_approx_equal(u_t[1000], 0.00017, significant=2)
np.testing.assert_approx_equal(u_t[-1], 0.000093, significant=2)
```

```python
plot = k3d.plot()
plot += k3d.mesh(mesh.points(), mesh.fv_indices(), attribute=init_u, color_map=k3d.matplotlib_color_maps.viridis)
plot += k3d.mesh(mesh.points() + np.array([-1.,0,0]), mesh.fv_indices(), attribute=u_t, color_map=k3d.matplotlib_color_maps.viridis)
plot
```

<!-- grade_id: cell-6f1953985c72acec, points: 1, grade: true, solution: false -->
```python
u_t_grad = compute_gradient_on_trimesh(mesh, u_t)
X = compute_gradient_of_distance_field(u_t_grad)
np.testing.assert_approx_equal(X[0,0], -0.59, significant=2)
np.testing.assert_approx_equal(X[0,1], -0.38, significant=2)
np.testing.assert_approx_equal(X[0,2], 0.71, significant=2)
np.testing.assert_approx_equal(X[-1,0], 0.06, significant=2)
np.testing.assert_approx_equal(X[-1,1], -0.89, significant=2)
np.testing.assert_approx_equal(X[-1,2], -0.46, significant=2)
np.testing.assert_approx_equal(X[1000,0], -0.96, significant=2)
np.testing.assert_approx_equal(X[1000,1], -0.23, significant=2)
np.testing.assert_approx_equal(X[1000,2], -0.19, significant=2)
```

```python
plot = k3d.plot()
plot += k3d.mesh(mesh.points(), mesh.fv_indices(), attribute=u_t, color_map=k3d.matplotlib_color_maps.viridis)
plot += k3d.vectors(mesh.points()[mesh.fv_indices()].mean(1), X*0.05, use_head=False)
plot
```

<!-- grade_id: cell-f904cbb4f55213b1, points: 1, grade: true, solution: false -->
```python
div = compute_divergence_on_trimesh(mesh, X)
phi = compute_distance_field(W, div)
phi -= phi.min() # solution is not unique up to constant factor. Therefore set minimum to zero
np.testing.assert_approx_equal(div[0], 0.0096, significant=2)
np.testing.assert_approx_equal(div[1000], 0.0056, significant=2)
np.testing.assert_approx_equal(phi[0], 1.41, significant=2)
np.testing.assert_approx_equal(phi[1000], 1.09, significant=2)
```

```python
k3d.mesh(mesh.points(), mesh.fv_indices(), attribute=div, color_map=k3d.matplotlib_color_maps.viridis)
```

```python
k3d.mesh(mesh.points(), mesh.fv_indices(), attribute=phi, color_map=k3d.matplotlib_color_maps.coolwarm)
```

```python
def plot_geodesics_heat(A,W,vertex):
    init_u = np.zeros(mesh.n_vertices())
    init_u[vertex] = 1.
    plot = k3d.plot()
    u_t = compute_heat_distribution(A, W, 1, init_u)
    u_t_grad = compute_gradient_on_trimesh(mesh, u_t)
    X = compute_gradient_of_distance_field(u_t_grad)
    div = compute_divergence_on_trimesh(mesh, X)
    phi = compute_distance_field(W, div)
    plot += k3d.mesh(mesh.points(), mesh.fv_indices(), attribute=phi, color_map=k3d.matplotlib_color_maps.coolwarm)
    plot.display()

interact(plot_geodesics_heat,
         A = fixed(A),
         W = fixed(W),
         vertex = range(mesh.n_vertices()),
         continous_update=False)
```

## Chamfer Distance
Compute the Chamfer distance with `chamfer_distance` for the two point clouds below. You can use `cdist` to compute the pairwise (Euclidean) distance matrix between the two point clouds.

```python
plot = k3d.plot()
plot += k3d.points(mesh.points(), point_size=0.02, color=0x0000ff)
plot += k3d.points(bob_mesh.points(), point_size=0.02, color=0xff0000)
plot
```

<!-- grade_id: cell-65959c4431b7ce3c, points: 0, grade: false, solution: true -->
```python
def chamfer_distance(pts_A, pts_B):
    ### BEGIN SOLUTION
    dists = cdist(pts_A, pts_B)
    return dists.min(1).mean() + dists.min(0).mean()
    ### END SOLUTION
```

<!-- grade_id: cell-670e5e2c5f459188, points: 1, grade: true, solution: false -->
```python
pts_A = mesh.points()
pts_B = bob_mesh.points()
d = chamfer_distance(pts_A, pts_B)
np.testing.assert_approx_equal(d, 0.51, significant=2)
```